# Welcome to this programming assignment
During week 1 we covered some basic numpy methods to perform basic calculations in Python. In this notebook you will be asked to complete some problems dealing with matrix and vector operations. There will be comments along the way to help you successfully complete this assignment. Good luck!

Check the [documentation](https://numpy.org/doc/1.17/reference/index.html) to find out more information about the library and methods it provides.

```python
!pip install numpy
import numpy as np
import numpy.linalg as LA
```

Output:
```
Requirement already satisfied: numpy in c:\users\home\miniconda3\envs\py-scripts\lib\site-packages (1.23.4)
```

## Task 1
Find the sum of all elements in an array - **ans1**. See [here](https://docs.scipy.org/doc/numpy/reference/generated/numpy.sum.html) for reference.

```python
A = np.array([[1,13,4,5, -6],[-9,4,12,0,2],[45,4,3,2,1], [3,-9,15,6,22],[3,5,0,-7,12]])
A
```

Output:
```
array([[ 1, 13,  4,  5, -6],
       [-9,  4, 12,  0,  2],
       [45,  4,  3,  2,  1],
       [ 3, -9, 15,  6, 22],
       [ 3,  5,  0, -7, 12]])
```

<!-- grade_id: cell-94f354b0c53df866, points: 0, grade: false, solution: true -->
```python
ans1 = np.sum(A)
```

<!-- grade_id: cell-293d329c86d14275, points: 1, grade: true, solution: false -->
```python
## GRADED PART, DO NOT CHANGE!
assert ans1 == 131
```

## Task 2
Calculate the difference between matrices m1 and m2 - **ans2**.

```python
m1 = np.array([[-7, -6,  3, -5],
       [ 7,  0,  4,  0],
       [ 6, -1,  1,  1]])
m2 = np.array([[  6,  -7,  -3,  -9],
       [ -8,   1,  -4,   2],
       [-10,   1,  -7,   9]])
```

<!-- grade_id: cell-3989f5b147e94477, points: 0, grade: false, solution: true -->
```python
ans2 = m1 - m2
```

<!-- grade_id: cell-a48c22ff6245a272, points: 1, grade: true, solution: false -->
```python
## GRADED PART, DO NOT CHANGE!
ans2 = ans2.flatten()
ans2 = np.sum(ans2)
assert ans2 == 32
```

## Task 3
Find the sum of all matrix elements whose absolute value is greater than 5 - **ans3**. You can use a $\textit{for}$ loop, but try using numpy methods (i.e. [this](https://docs.scipy.org/doc/numpy/reference/generated/numpy.where.html)), since the numpy library is very fast and effective in manipulating data.

```python
b = np.array([[ -6,  -7,  -5,   0],
       [-15,   2, -13,  -5],
       [  8,   4,  10,   8],
       [ -1,  -9,   7,  -4]])
```

<!-- grade_id: cell-41eb95f9a6606e5f, points: 0, grade: false, solution: true -->
```python
ans3 = np.sum(b, where=np.abs(b)>5)
```

<!-- grade_id: cell-d2759db18abc6992, points: 1, grade: true, solution: false -->
```python
assert ans3 == -17
```

## Task 4
Find the product of two matrices - **ans4**. Remember that it is better to use $\textit{matmul, or @,}$ for $\textbf{larger dimension}$ matrices.

```python
n = np.array([[ -4,  -1,   3,  -7,   6,  -9,   5,  -9,  -7,   0],
       [-12,   3, -12,  -7,   2,  -4,   8,  13,   9,   3],
       [ -2,  -1, -15, -10,   7,   5,   7, -13,  -9,   6],
       [  7,  11,  -5,  -6,  -9,  -1,   3,  10, -15,  -6],
       [ -3,  -4,   4, -13,   4, -13,  10,  -6,  -8,  13],
       [  7, -13,  -4,   3,   4,  13, -13,   4,  -9, -10],
       [ -5,  13,  -9, -15,  11,   2,   7, -14, -14,   2],
       [ -4,  -6,  12,  10,  13,  -3,  12, -15,   5,  -2],
       [ -1,   8,   3,   7,  11,  -3,   6,  12,  12, -13],
       [ 12,  11,  12,  11,   1, -14,   2,   6,  -9,  -3]])
z = np.array([[-15,   5, -15,  10,  -3,  12,  -6, -15,   6,   3],
       [ 12,  -4,   6, -11,   5,   3,   3,  12, -15,  -6],
       [-11, -11, -10,   9,  -8, -14, -14,   7,  -2,  11],
       [  5,  -4,  11,   0,   8,  -9, -15, -11,   7, -10],
       [ -6,  -9,  -3, -13, -11,  -9,   7,  -7,   9,   4],
       [ 10,  12,   4,  -7,   1,   0,  -1,  11,  -4,  11],
       [-14,   2,  -8,  11,  -9,   5, -14,  10,  -6,  -9],
       [ -5,  10, -10,  -1,  -9, -12,  12,  14,   4,   2],
       [ -1, -13,  10,  11,  12, -14,   9,  12,  12,  -5],
       [  2,   6, -13,  -5,  -2,  -2,   0,  14,   7,  14]])
```

<!-- grade_id: cell-cb84fbb30694589c, points: 0, grade: false, solution: true -->
```python
ans4 = np.matmul(n,z)
```

<!-- grade_id: cell-26d17ea398c5775a, points: 1, grade: true, solution: false -->
```python
## GRADED PART, DO NOT CHANGE!
ans4 = ans4.flatten()
ans4 = np.sum(ans4)
assert ans4 == -5004
```

## Task 5
Create an identity matrix of the necessary size and subtract it from A.  Save the result to a variable **ans5**

```python
A = np.array([[  7,   8, -15,   8,  11,   7],
       [ -5,  -9,  13,  -3, -15, -11],
       [ -3,   0,  10,  -5,   3,  -8],
       [ 10,   1,   7, -11,  10,  -3],
       [ 11,   9,  -6, -14,  10, -11],
       [  6,  -4,  -4,  -9,  -6,  -7]])
A
```

Output:
```
array([[  7,   8, -15,   8,  11,   7],
       [ -5,  -9,  13,  -3, -15, -11],
       [ -3,   0,  10,  -5,   3,  -8],
       [ 10,   1,   7, -11,  10,  -3],
       [ 11,   9,  -6, -14,  10, -11],
       [  6,  -4,  -4,  -9,  -6,  -7]])
```

<!-- grade_id: cell-f83875bf4ad464c3, points: 0, grade: false, solution: true -->
```python
ident_A = np.identity(A.shape[0])
ans5 = A - ident_A
```

<!-- grade_id: cell-db5252681a67d532, points: 1, grade: true, solution: false -->
```python
## GRADED PART, DO NOT CHANGE!
ans5 = ans5.flatten()
ans5 = np.sum(ans5)
assert ans5 == -24
```

## Task 6
Find the solution to system Ax = b, use matrix inverse. Save the result to a variable **x**

```python
A = np.array([[ 0, -2,  0, -2],
       [-2,  0, -4,  4],
       [-1, -5, -2,  2],
       [ 3, 1, 2,  1]])
```

```python
b = np.array([-4,  1,  4,  2])
```

<!-- grade_id: cell-d3207e36153af09a, points: 0, grade: false, solution: true -->
```python
x = LA.inv(A) * b
```

<!-- grade_id: cell-087b4a594615c112, points: 1, grade: true, solution: false -->
```python
x = np.round(np.sum(x),3)
assert x == 3.225
```

## Task 7
Find the PLU decomposition and assert that by multiplying the matrices in the decomposition you get the original matrix back. Submit the matrix U as the answer.

```python
from scipy.linalg import lu
```

```python
m = np.array([[ 12, -13,   2,   2, -11],
       [  1,   4,  11,  11,  -5],
       [ 13,  -5,   9,  12,  -2],
       [  5, -13,  12,   3,   5],
       [ -9,  11,  -1,  -6,  -5]])
m
```

Output:
```
array([[ 12, -13,   2,   2, -11],
       [  1,   4,  11,  11,  -5],
       [ 13,  -5,   9,  12,  -2],
       [  5, -13,  12,   3,   5],
       [ -9,  11,  -1,  -6,  -5]])
```

<!-- grade_id: cell-00a31881b2ccb672, points: 0, grade: false, solution: true -->
```python
P, L, U = lu(m)
assert np.allclose(np.dot(np.dot(P, L), U), m)
```

<!-- grade_id: cell-cee5258f3970597b, points: 1, grade: true, solution: false -->
```python
## GRADED PART, DO NOT CHANGE!
U = U.flatten()
U = np.sum(U)
U = round(U)
assert U == 26
```

```python

```